<?xml version="1.0" encoding="UTF-8"?>
<quiz>

  <question type="coderunner">
    <name>
      <text>Demo simple Octave function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Write a Matlab function <i>sqr(x)</i>&nbsp;that returns the square of the parameter x.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <coderunnertype>octave_function</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime></penaltyregime>
    <showsource>0</showsource>
    <answerboxlines>8</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <useace>1</useace>
    <resultcolumns></resultcolumns>
    <answer>function square = sqr(x)
    square = x * x;
end</answer>
    <combinatortemplate></combinatortemplate>
    <testsplitterre></testsplitterre>
    <enablecombinator></enablecombinator>
    <pertesttemplate></pertesttemplate>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <testcases>
      <testcase useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>disp(sqr(7))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>49</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>disp(sqr(-11))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>121</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>disp(sqr(0))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>0</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>disp(sqr(-101))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>10201</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase useasexample="0" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>disp(sqr(20))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>400</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>
